# Basic list exercises
This notebook is a quick check-in to see how you are progressing on learning Python.  You will be required to complete 5
functions based on the descriptions.

Please download this file to your local machine and upload and submit when you are finished. You will recieve feedback on your results immediately. Good luck!

## Honour Code

I **YOUR NAME, YOUR SURNAME**, confirm - by submitting this document - that the solutions in this notebook are a result of my own work and that I abide by the EDSA honour code (https://drive.google.com/file/d/1QDCjGZJ8-FmJE3bZdIQNwnJyQKPhHZBn/view?usp=sharing).  

Non-compliance with the honour code constitutes a material breach of contract.

## match_ends
Given a list of strings, return the count of the number of strings where the string length is 2 or more and the first and last chars of the string are the same.<br>
Note: python does not have a ++ operator, but += works.

<!-- grade_id: match_ends, points: 0, grade: false, solution: true -->
```python
def match_ends(words):
    # YOUR CODE HERE
    raise NotImplementedError()
    return None
```

Ensure that your method passes the following assertions (against which it will be tested):

<!-- grade_id: correct_match_ends, points: 1, grade: true, solution: false -->
```python
assert match_ends(['abcdea', 'abcxa']) == 2 
assert match_ends(['tallest', 'lovely', 'julsfsj']) == 2 
```

<!-- grade_id: cell-89048d0c01324341, points: 1, grade: true, solution: false -->
```python
match_ends(['abcdea', 'abcde'])
```

# front_x
Given a list of strings, return a list with the strings in sorted order, except group all the strings that begin with 'x' first.<br>
e.g. ['mix', 'xyz', 'apple', 'xanadu', 'aardvark'] yields<br>
`['xanadu', 'xyz', 'aardvark', 'apple', 'mix']`<br>
Hint: this can be done by making 2 lists and sorting each of them before combining them.

<!-- grade_id: front_x, points: 0, grade: false, solution: true -->
```python
def front_x(words):
    # YOUR CODE HERE
    raise NotImplementedError()
    return None
```

Be sure to pass the following test cases:

<!-- grade_id: correct_front_x, points: 2, grade: true, solution: false -->
```python
assert front_x(['mix', 'xyz', 'apple', 'xanadu', 'aardvark']) == ['xanadu', 'xyz', 'aardvark', 'apple', 'mix']
assert front_x(['netowrk', 'artist', 'xamarian', 'king', 'cat']) == ['xamarian', 'artist', 'cat', 'king', 'netowrk'] 
```

```python
front_x(['mix', 'xyz', 'apple', 'xanadu', 'aardvark'])
```

## sort_last
Given a list of non-empty tuples, return a list sorted in increasing order by the last element in each tuple.<br>
e.g. [(1, 7), (1, 3), (3, 4, 5), (2, 2)] yields<br>
`[(2, 2), (1, 3), (3, 4, 5), (1, 7)]`<br>
Hint: use a custom key= function to extract the last element form each tuple.

<!-- grade_id: sort_last, points: 0, grade: false, solution: true -->
```python
def sort_last(tuples):
    # YOUR CODE HERE
    raise NotImplementedError()
    return None
```

<!-- grade_id: correct_sort_last, points: 2, grade: true, solution: false -->
```python
assert sort_last([(1, 7), (1, 3), (3, 4, 5), (2, 2)]) == [(2, 2), (1, 3), (3, 4, 5), (1, 7)]
assert sort_last([(3, 1, 9), (2, 6), (5, 7, 1)]) == [(5, 7, 1), (2, 6), (3, 1, 9)] 
```

```python
sort_last([(1, 7), (1, 3), (3, 4, 5), (2, 2)])
```

## remove adjacent 
Given a list of numbers, return a list where all adjacent == elements have been reduced to a single element, so [1, 2, 2, 3] returns [1, 2, 3]. You may create a new list or modify the passed in list.

<!-- grade_id: remove_adjacent, points: 0, grade: false, solution: true -->
```python
def remove_adjacent(nums):
    # YOUR CODE HERE
    raise NotImplementedError()
    return None
```

Be sure to preserve the order (see below) and ensure that your code passes the assertions below:

<!-- grade_id: correct_remove_adjacent, points: 1, grade: true, solution: false -->
```python
assert remove_adjacent([1, 2, 2, 3]) == [1, 2, 3]
assert remove_adjacent([1, 2, 2, 2, 3, 3]) == [1, 2, 3]
assert remove_adjacent([27, -5, 13, 13, 0]) == [27, -5, 13, 0]
```

<!-- grade_id: bonus_correct_remove_adjacent, points: 1, grade: true, solution: false -->
```python
remove_adjacent([1, 2, 2, 3])
```

## Linear Merge
Given two lists sorted in increasing order, create and return a merged list of all the elements in sorted order. You may modify the passed in lists. <br>
Ideally, the solution should work in "linear" time (dont worry if you don't know what this is, we will be covering it next week), making a single pass of both lists.

<!-- grade_id: linear_merge, points: 0, grade: false, solution: true -->
```python
def linear_merge(list1, list2):
    # YOUR CODE HERE
    raise NotImplementedError()
    return None
```

<!-- grade_id: correct_linear_merge, points: 2, grade: true, solution: false -->
```python
assert linear_merge(['aa', 'xx', 'zz'], ['bb', 'cc']) == ['aa', 'bb', 'cc', 'xx', 'zz']
assert linear_merge(['aab', 'aad', 'aaf'], ['aac', 'aae']) == ['aab', 'aac', 'aad', 'aae', 'aaf']
```

```python
linear_merge(['aa', 'xx', 'zz'], ['bb', 'cc'])
```

# References

This tutorial was created using the wonderful [Google Python Developers](https://developers.google.com/edu/python/) course, and has been modified under the [Creative Common's licence 2.5](https://creativecommons.org/licenses/by/2.5/) 

Copyright 2010 Google Inc.
Licensed under the Apache License, Version 2.0
http://www.apache.org/licenses/LICENSE-2.0

Google's Python Class
http://code.google.com/edu/languages/google-python-class/